Ubiquitin ligase and signalling hub MYCBP2 is required for efficient EPHB2 tyrosine kinase receptor function

Author Response

The following is the authors’ response to the previous reviews

The revised manuscript is much improved - many unclear points are now better explained. However, in our opinion, some issues could still be significantly improved.

1. Statistics: none of us are experts in statistics but several things remain questionable in our opinion and if it were our study, we would consult with an expert:

a) while we understand the authors note about N-chasing and p-hacking, we wonder how the number of N's was premeditated before obtaining the results. Why in 4M an N of 3 is sufficient while in 3E the N is >20 (and not mentioned). At the very least, we think it would be wise to be cautious when stating something as not-significant when it is clear (as in 4M) that the likelihood of it actually being statistically significant is quite large.

b) In most analyses, the data is not only normalized by actin or some other measure but also to the first (i.e left side on the graph) condition, resulting in identical data points that equal '1' (in Figure 4 alone - C; I; K; M; and O) - while this might be scientifically sound, it should be mentioned (the specific normalization) and also note that this technique shadows any real variance that exists in the original data in this condition. consider exploring techniques to overcome this issue.

c) In 3C, - if we understand the experiment, you want to convince us that the DIFFERENCE between eB2-FC compared to FC is larger in the control compared to the experiment. We are not absolutely sure that the statistical tools employed here are sufficient - which is why we would consult an expert.

A) We are aware that many studies do not consistently quantify such experiments. For example, there are essentially no published examples of the signalling timelines of EphB2 receptors as in Fig. 5. By striving to quantifying such biochemical effects, an unquantified experiment stands out, and so perhaps we were too strict by trying to quantify as many experiments as possible, resulting in low n’s for some of them. We acknowledge that additional experiments on EPHB1 protein stability may reach significance. We have adjusted our text on line 332-335 to point to this interesting trend, and slightly changed the conclusion to this section. Similarly, we commented on similar trends when describing Figs. 1E and 4G on lines 901 and 952.

B) For the Western blot band intensity normalisation, we believe that our method is scientifically sound. Normally, when the replicate samples are loaded on one gel and blotted on the same membrane, the experimenter only needs to normalise the target band intensity to its cognate loading control band intensity for quantitation. However, we usually have a large number of samples from multiple experiments, carried out on different dates. For example, in Fig. 4B,C there are 7 biological replicates collected from 7 experiments and in Fig. 4D there are 10 protein samples. It is not possible for us to run all samples on the same gel. In addition, due to the combined effects of variance in transfer efficiency, the potency of antibodies, detection efficiency and the developing time for each blot, it is practically impossible to generate similar band intensity for each batch. Thus, we use normalisation of test bands to the loading control for individual experiments, and this analysis method is widely accepted by reputable journals with a focus on biochemical experiments (for example: PMID 37695914: Fig. 3 A,B,C; PMID 36282215: Fig. 3 B,C,D,E; PMID 33843588: Fig. 3 C,D,E,F,G,H). Since the value of the first sample on the plot is 1, which is a hypothetical value and does not meet the parametric test requirement, we performed one-sample t-test for statistics when other samples are compared with the first sample (PMID 35243233 Fig. 6 A,B,C,D; https://www.graphpad.com/quickcalcs/oneSampleT1/, “A one sample t-test compares the mean with a hypothetical value. In most cases, the hypothetical value comes from theory. For example, if you express your data as 'percent of control', you can test whether the average differs significantly from 100.”). Thus, we believe that our normalisation and statistical methods are both correct with a large number of precedents.

C) This comment refers to the cell collapse experiment shown in Fig. 3C for which the data are plotted in Fig. 3D. We stand by the statistical method used. There are two groups of cells (CTRLCRISPR and MYCBP2 CRISPR) and two treatments for each cell group (Fc control and eB2), thus we should use two-way ANOVA. Since we compared the cell retraction effects of Fc and eB2 on the two groups of cells, Sidak post hoc comparison is the right method to avoid errors introduced by multiple comparisons. Here is an example of an eLife article that used the same statistical method for similar comparisons: PMID 37830910, Fig. 1 H,I. To make the comparison easier, we grouped the experiments by cell type (CTRLCRISPR and MYCBP2 CRISPR) as opposed to by treatment. Below, the old version is on the right, and the new version is on the left. The conclusion is that eB2 induces less cell collapse in cells depleted of MYCBP2, when compared to the control cells. However, eB2 is still able to collapse cells lacking MYCBP2.

Revisiting these data, we noticed an error introduced when CC compiled the data used to generate Fig. 3D. The data were acquired from nine biological replicates per condition. CC used a mix of two methods for cell collapse rate calculation: the first method involved the sum of collapsed cells and all cells from multiple regions of one coverslip (biological replicate). The second method involved computing a collapse rate in each region which then was used to calculate the average collapse rate for the entire coverslip (technical replicate). Given the small cell numbers due to sparse culture conditions, we believe that the first method is a more conservative approach. We hence re-plotted all replicate data using the first method. This resulted in slightly different % collapse and p values. These were changed accordingly in the text and plot and do not affect the conclusion of this experiment.

1. thanks for the clarification that the interaction between the extracellular domain of EPHB2 and MYCBP2 might not occur directly - however, unless we missed this it was not clearly stated in the text. It is an important point and also a cool direction for the future - to find the elusive co-receptor that actually helps EPHB2 and MYCBP2 form a complex.

We now also refer to this in the results section on line 215.

“Since EPHB2 is a transmembrane protein and MYCBP2 is localised in the cytosol, these experiments suggest that the interaction between the extracellular domain of EPHB2 and MYCBP2 might be indirect and mediated by other unknown transmembrane proteins.”

1. The Hela CRISPR cell line is better explained in the response letter but still not sufficiently explained in the text for a non-expert reader. If the authors want any reader to comprehend this, we would strongly recommend adding a scheme.

We now include a schematic outlining the CRISPR cell generation as Fig. 3A and its description on line 926.

1. To clarify some of our previous (and persisting) concerns about Figure 3D/E - it is true that a reduction in 25% of cell size is dramatic. But (if we understand correctly) your claim is that a reduction in 22% (this is a guess, as the actual numbers are not supplies) is significantly less than 25%. Even if it is, statistically speaking, significant, what is the physiological relevance of this very slight effect? In this experiment, the N was quite large, and we wonder if the images in D are representative - it would be nice to label the data points in E to highlight which images you used.

We now mention the average cell area contraction measurements in the legend to Fig. 3F on line 935. We also tracked down the individual cells shown in Fig. 3E and they are now labelled as data points in blue in Fig. 3F. HeLa cell collapse is a simplified model of EPHB2 function and we do not know whether the difference between the behaviour of CTRLCRISPR and MYCBP2 CRISPR cells is physiologically significant and thus we prefer not to speculate on this.

1. Figure 3F and other stripe assays - In the end, it is your choice how to quantify. We believe that quantifying area of overlap is a more informative and objective measurement that might actually benefit your analyses. That said, if you do keep the quantification as it is now, you have to define the threshold of what you mean by "cell/s (or an axon in 7A, where it is even more complicated as are you eluding to primary, secondary, or even smaller branches) are RESIDING within the stripe". Is 1% overlap sufficient or do you need 10 or 50% overlap?

We now added this statement to the methods on line 745: “A cell was considered to be on an ephrin-B2 stripe when more than 50% of its nucleus was located on that stripe”. For chick explant stripe assay, when measuring the length of an axon on a stripe, we only measured the main axons originated from the explants.

For explant/stripe experiments in Fig. 7 AB, we now use the term “GFP-expressing neurite” rather than “branch”. This was already present in the results of the previous version, but the methods and legend needed to be brought up to date (lines 786 and 1008. We think that “branch” was a confusing term that was supposed to mean the same thing as “neurite” but came across as some indication of branching. We do not know whether the GFP+ neurites were primary or secondary extensions of explants, or in fact, whether some of them contained more than one axon. We also adjusted the method to reflect the fact that some stripes were used in conjunction with a single explant and added a reference to a previous study extensively using this method (Poliak et al., 2015) on line 778.

1. We still don't get the link to the lysosomal degradation. Your data suggests that in your cells EPHB2 is primarily degraded by the lysosomal pathway and not proteasome. Any statement about MYCBP2 is not strongly supported by the data, in our opinion - Unless you develop some statistical measurement that shows that the effect of BafA1 is statistically different in MYCBP2 cells than in control cells. Currently, this is not the case and the link is therefore not warranted in our opinion.

We generated a new version of Fig. 4K with average increase in EPHB2 levels in the presence of BafA1 and CoQ, compared to DMSO treated controls (see below). BafA1 and CoQ restored EPHB2 protein levels by 19% and 14% respectively in CtrlCRISPR cells, while the inhibitors restored EPHB2 protein levels by 40% and 35% respectively in MYCBP2 CRISPR cells.

For each of the 4 replicates, the increase in EPHB2 levels by BafA1 compared to DMSO is as follows:

These values are not significantly different between CtrlCRISPR cells versus MYCBP2 CRISPR cells (p= 0.08, student’s t test). Similarly for the CoQ experiment. We now temper our conclusion for this experiment: Although the difference in percentage increase between CTRLCRISPR cells and MYCBP2CRISPR cells is not significant, this trend raises the possibility that the loss of MYCBP2 promotes EPHB2 receptor degradation through the lysosomal pathway (line 319). We also adjusted the section title (line 306).

1. While the C. elegans part is now MUCH better explained - we are not sure we understand the additional insight. The fact that vab-1 and glo4 double mutants are additive as are vab1 and fsn1, suggest they act in parallel (if the mutants are NULL, and not if they are hypomorphs, if one wants to be accurate) - how this relates to your story is unclear. The vab1/rpm1 double mutant is still uninformative and incomplete. rpm1 phenotype is so severe that nothing would make it more severe. We read the Jin paper that the authors directed to - nothing makes the rpm1 phenotype more severe. Yes, some DOWNSTREAM elements make the rpm1 phenotype LESS severe - this is not something you were testing, to the best of our knowledge. Rather, you wanted to see if rpm1 mutant resulted in stabilization of vab1 and thus suppression of vab1 phenotype - we are just not sure the system is amenable to test (actually reject) your hypothesis that Vab1 is degraded by rpm1. Also, assuming we are talking about NULLs, the fact that the rpm1 phenotype is WAY stronger than the vab1 mutant, suggests that rpm1 functions via multiple routes, adding even more complexity to the system. Given these results, despite the much improved clarity, we are still not sure that the worm data adds new insight, rather than potentially confusing the reader.

We realise that the genetic interactions between vab-1 and the RPM-1/MYCBP2 signalling network are complicated. However, we insist on keeping the data for the sake of its availability for future studies and completeness. We also think it is important for readers and the community to see these data, even if the authors and reviewers are not entirely in agreement about the importance/interpretation of experimental outcomes. It is our hope that the community will examine the results and draw their own conclusions.

A few points of clarification:

The C. elegans experiments were designed to test genetically if the vertebrate interactions between EPHB2 and MYCBP2 and its signalling network are conserved. We studied two kinds of interactions: (1) between vab-1 and RPM-1/MYCBP2 downstream proteins (GLO-4 and FSN-1) and (2) between vab-1 and rpm-1. For these studies, we used null alleles for vab-1, glo-4 and fsn-1 which is now noted on lines 440, 453, 475 and 859. Our findings are consistent with the VAB-1 Ephrin receptor functioning in parallel to known RPM-1 binding proteins. This is further supported by new data: vab-1; fsn-1 double mutants showed enhanced incidence of axon overextension defects using a second transgenic background, zdIs5 (Pmec-4::GFP), to visualize axon termination (Fig. 8F).

This second transgenic background also allowed us to generate new data to address your concerns about phenotypic saturation in rpm-1 mutants. To do this, we used the zdIs5 (Pmec4::GFP) genetic background, in which axon termination defects are not saturated in rpm-1 mutants (Fig. 8F) because they can be enhanced by other mutants such as cdc-42 and unc-33 (Fig. 7C, D, in Borgen et al. Development 144, 4658–4672 (2017), PMID 29084805). In this new background, we found that vab-1 loss of function fails to enhance the incidence of severe “hook” defects in rpm-1 mutants which is an indication that the two genes function in the same pathway. Importantly, prior studies in this background, also showed that mutants in the RPM-1 signalling network (e.g. fsn-1, glo-4 and ppm-2) do not enhance the incidence of severe “hook” defects as double mutants with rpm-1 compared to rpm-1 single mutants (Fig. 7B, ibid.).

To reflect these ideas more clearly, we revised the Results section pertaining to C. elegans genetics (starting on line 418) and tempered our discussion (lines 517). Basically, this section now says that we studied genetic interactions between vab-1 and the RPM-1/MYCBP2 signalling network. From these experiments we conclude that: (1) The enhancement of overextension defects in vab-1; glo-4 and vab-1; fsn-1 double mutants compared to single mutants indicates that VAB-1/EPHR functions in parallel to known RPM-1 binding proteins to facilitate axon termination, and (2) Since the vab-1; rpm-1 double mutants do not display an increased frequency or severity of overextension defects compared to rpm-1 single mutants, VAB-1 /EPHR functions in the same genetic pathway as RPM-1/MYCBP2.

The new genetic data included in this version were generated by Karla J. Opperman who is now included as a co-author.

Further corrections:

Because of the errors associated with quantifications in Fig. 3D (see above), we reviewed other quantification methodologies and noticed another discrepancy that required a correction. In the hippocampal neuron growth cone collapse assay shown in the previous version of Fig. 7 D (left), the growth cones were classified into three groups: 1, fully collapsed; 2, hard to tell, but not fully collapsed; 3, fan-shape cones. Two different quantifications were performed as follows: (1), number of fully collapsed cones divided by the numbers of all growth cones; (2), number of fully collapsed cones divided by [number of fully collapsed cones + fan-shape cones]. CC erroneously used the second method to generate Fig. 7D.

We think that the first method is more appropriate. Furthermore, since n=5 for the Fc and eB1-Fc conditions, but n=3 for the eB2-Fc condition, we decided to omit it. The final plot for figure 7D is the following:

Our conclusion still stands that exogenous FBD1 WT overexpression impaired the growth cone collapse mediated by EphB.

Changes from the earlier version to the revised version

Abstract

Abstract

Eph receptor tyrosine kinases participate in a variety of normal and pathogenic processes during development and throughout adulthood. This versatility is likely facilitated by the ability of Eph receptors to signal through diverse cellular signalling pathways: primarily by controlling cytoskeletal dynamics, but also by regulating cellular growth, proliferation, and survival. Despite many proteins linked to these signalling pathways interacting with Eph receptors, the specific mechanisms behind such links and their coordination remain to be elucidated. In a proteomics screen for novel EPHB2 multi-effector proteins, we identified human MYC binding protein 2 (MYCBP2 or PAM or Phr1). MYCBP2 is a large signalling hub involved in diverse processes such as neuronal connectivity, synaptic growth, cell division, neuronal survival, and protein ubiquitination. Our biochemical experiments demonstrate that the formation of a complex containing EPHB2 and MYCBP2 is facilitated by FBXO45, a protein known to select substrates for MYCBP2 ubiquitin ligase activity. Formation of the MYCBP2-EPHB2 complex does not require EPHB2 tyrosine kinase activity and is destabilised by binding of ephrin-B ligands, suggesting that the MYCBP2-EPHB2 association is a prelude to EPHB2 signalling. Paradoxically, the loss of MYCBP2 results in increased ubiquitination of EPHB2 and a decrease of its protein levels suggesting that MYCBP2 stabilises EPHB2. Commensurate with this effect, our cellular experiments reveal that MYCBP2 is essential for efficient EPHB2 signalling responses in cell lines and primary neurons. Finally, our genetic studies in C. elegans provide in vivo evidence that the ephrin receptor VAB-1 displays genetic interactions with known MYCBP2 binding proteins. Together, our results align with the similarity of neurodevelopmental phenotypes caused by MYCBP2 and EPHB2 loss of function, and couple EPHB2 to a signaling effector that controls diverse cellular functions.

Introduction

Eph receptor tyrosine kinases and their membrane-tethered ligands, the ephrins, elicit short distance cell-cell signals that regulate many biological processes. Ephrin-Eph signalling primarily impacts the cytoskeleton with the immobilization of highly dynamic axonal growth cones being a classic example. Other processes that involve changes in transcription, growth, and survival such as angiogenesis, synaptic plasticity, stem cell fate, tumorigenesis and neurodegeneration also involve the Eph/ephrin system. Many proteins are postulated to couple Eph receptors to different intracellular effectors, but the molecular logic of this diversity remains fragmented (Kania and Klein, 2016;Bush, 2022).

EphB subfamily members preferentially bind transmembrane ephrin-Bs and although both molecules participate in bidirectional signalling, ephrin-B activation of EphB signalling cascades is more thoroughly studied (Gale et al., 1996;Mellitzer et al., 1999). To elicit robust Eph receptor forward signalling, ephrins multimerise in signalling clusters by intercalating with Ephs on a signal-recipient cell with array size correlating with signal amplitude (Kullander et al., 2001;Schaupp et al., 2014). Signalling initiation involves the activation of the receptor tyrosine kinase and phosphorylation of tyrosines proximal to the EphB transmembrane domain (Soskis et al., 2012;Binns et al., 2000). Eph-evoked cytoskeletal effects such as cell contraction and growth cone collapse result from changes in small GTPase activity modulated by EphB Guanine nucleotide exchange factors (GEFs) and GTPase activating proteins (GAPs) (Margolis et al., 2010;Shi et al., 2007). Eph signalling has also been linked to fundamental cellular pathways such as the Ras-MAPK pathway, mTOR-regulated protein synthesis, cell division and survival (Bush and Soriano, 2010;Nie et al., 2010;Fawal et al., 2018;Genander et al., 2009;Depaepe et al., 2005). Eph forward signalling eventually leads to their internalisation and either recycling or degradation via endosome/lysosome and ubiquitination/proteasome pathways (Zimmer et al., 2003;Okumura et al., 2017). While identification of a growing number of proteins interacting with Eph receptors has moved the field forward, we have yet to clarify the question of how Eph receptors activate various fundamental cellular processes, often within the same cell type.

Myc-binding protein 2 (MYCBP2), also known as Protein Associated with Myc (PAM) and Highwire, RPM-1, or Phr1 in different species, is a large signalling hub that regulates cytoskeletal dynamics, neuronal development, and axonal degeneration (Guo et al., 1998;Grill et al., 2016;Virdee, 2022). It has an atypical RING ubiquitin ligase activity that inhibits the p38/MAP and JNK kinase pathways thereby regulating cytoskeletal dynamics underlying axonal development and synaptic growth (Nakata et al., 2005;Collins et al., 2006;Pao et al., 2018;Wan et al., 2000;Lewcock et al., 2007;Borgen et al., 2017). MYCBP2 further regulates the Tuberin Sclerosis Complex linked to cell growth (Han et al., 2012), initiation of autophagy via ULK (Crawley et al., 2019) and NMNAT2-regulated neuronal survival and axonal degeneration (Babetto et al., 2013;Xiong et al., 2012). Biochemical mapping has shown that human MYCBP2 and its C. elegans ortholog RPM-1 rely upon the FBD1 domain to bind the F-box protein FBXO45 that acts as a ubiquitination substrate selector (Desbois et al., 2018;Sharma et al., 2014). Intriguingly, the neurodevelopmental phenotypes caused by MYCBP2 and EPHB2 loss of function are similar, raising the possibility that these two molecules could function in the same pathway (Henkemeyer et al., 1996;Lewcock et al., 2007;Dalva et al., 2000). Nonetheless, a biochemical or genetic interaction between MYCBP2 and EPHB2 has not been demonstrated in any system.

To shed light on how EphB receptors fulfil their multitude of functions, we used mass spectrometry (MS)-based proteomics to identify multi-effector proteins that bind EPHB2. One of our proteomic hits was MYCBP2, which we demonstrated forms a complex with EPHB2 using a combination of biochemical and cellular assays. Furthermore, we show that this interaction is required for efficient EPHB2 signalling in cell lines and primary neurons. Consistent with these findings, we observed in vivo genetic interactions in C. elegans between the Eph receptor, VAB-1, and known RPM-1 binding proteins. Our collective results indicate that the relationship between EPHB2 and MYCBP2 does not appear to involve the ubiquitin ligase activity of MYCBP2, and raise the possibility that MYCBP2 links EPHB2 to diverse fundamental cellular functions.

Results

Proteomics identifies MYCBP2 as a putative EPHB2-interacting protein

To understand the molecular logic underlying EPHB2 signalling diversity, we performed affinity purification coupled to MS (AP-MS) in order to identify EPHB2-interacting proteins, and prioritised those known to be signalling hubs. We used a stable HeLa cell line with tetracycline-inducible expression of BirA-linked EPHB2-FLAG, that we previously used to study EPHB2 signalling (Lahaie et al., 2019). To identify ephrin ligand-dependent EPHB2 protein complexes, we stimulated EPHB2-FLAG-overexpressing cells with pre-clustered Fc control or ephrinB2-Fc (eB2-Fc). We then harvested and lysed the cells, performed anti-FLAG immunoprecipitation, and used mass spectrometry to identify EPHB2 protein complexes (Fig. 1A).

MS-proteomics and biochemistry in HeLa cells identifies MYCBP2 as EPHB2 binding protein.

Earlier version: A) Schematic of EPHB2 affinity purification coupled to mass spectrometry (AP-MS) workflow. B) Scatter plot of AP-MS data showing known and putative EPHB2 binding proteins, including MYCBP2. Y axis and X axes represent the average spectral counts of the identified protein hits in the EPHB2 protein complexes from cells stimulated with Fc control or ephrin-B2 (eB2-Fc), respectively. C) In HEK 293T cells, endogenous MYCBP2 is pulled down by transiently overexpressed EPHB2-FLAG but not by EPHA3-FLAG. D) In EPHB2-FLAG stable HeLa cell line, ephrin-B stimulation reduces the interaction between MYCBP2 and EPHB2. E) Quantification of MYCBP2-EPHB2 association intensity after Fc, ephrin-B1 (eB2-Fc) or ephrin-B2 (eB2-Fc) treatment (eB1-Fc, p=0.1365; eB2-Fc, p=0.0002; one-sample t-test). Error bars represent standard deviation (SD). F) Representative image of MYCBP2 pull down with anti-EPHB2 or IgG control antibodies from rat cortical neurons. Asterisk indicates MYCBP2. G) Representative images from western blot analysis of endogenous MYCBP2 following IP of GFP-EPHB2 wild-type (WT) or its kinase dead (KD) counterpart.

Revised version: A) Schematic of EPHB2 affinity purification coupled to mass spectrometry (AP-MS) workflow. B) Scatter plot of AP-MS data showing known and putative EPHB2 binding proteins, including MYCBP2. Y axis and X axes represent the average spectral counts of the identified protein hits in the EPHB2 protein complexes from cells stimulated with Fc control or ephrin-B2 (eB2-Fc), respectively. C) In HEK 293T cells, endogenous MYCBP2 is pulled down by transiently overexpressed EPHB2-FLAG but not by EPHA3-FLAG. D) In EPHB2-FLAG stable HeLa cell line, ephrin-B stimulation reduces the interaction between MYCBP2 and EPHB2. E) Quantification of MYCBP2-EPHB2 association intensity after Fc, ephrin-B1 (eB2-Fc) or ephrin-B2 (eB2-Fc) treatment (eB1-Fc, p=0.1365; eB2-Fc, p=0.0002; one-sample t-test). EPHB2-MYCBP2 interaction reduction evoked by eB1-Fc is not statistically significant, probably because of high experimental variability which could be biologically significant. Error bars represent standard deviation (SD). F) Representative image of MYCBP2 pull down with anti-EPHB2 or IgG control antibodies from rat cortical neurons. Asterisk indicates MYCBP2. G) Representative images from western blot analysis of endogenous MYCBP2 following IP of GFP-EPHB2 wild-type (WT) or its kinase dead (KD) counterpart.

To identify EPHB2-specific interactions and remove background contaminants, we performed Significance Analysis of INTeractome (SAINTexpress analysis (Teo et al., 2014)) using our EPHB2-related controls (Lahaie et al., 2019). To better visualize the changes in putative EPHB2 binding partners, we compared the average spectral counts of the identified proteins using ProHits-viz tool (Knight et al., 2017). Comparison of Fc and ephrin-B2-treated samples did not yield any significant differences. Since we applied the ligand and collected the samples on a time scale comparable to known EphB signalling dynamics, this limitation could be potentially due to autoactivation of Eph receptors when they are overexpressed, obscuring some ligand-dependent effects (Lackmann et al., 1998). However, the resulting scatter plot confirmed the presence of several known EPHB2 interactors, such as FYN and YES1 Src kinases and members of the Eph receptor family (Fig. 1B) (Banerjee et al., 2022). One of the most prominent, novel hits was the E3 ubiquitin ligase and signalling hub protein MYCBP2, and its binding partner FBXO45. FBXO45 was previously identified as a putative EPHB2 interacting protein in large-scale, cell-based interactome studies (Huttlin et al., 2021;Salokas et al., 2022).

Biochemical validation of MYCBP2 binding to EPHB2

To confirm that MYCBP2 can indeed form a molecular complex with EPHB2, we tested whether endogenous MYCBP2 co-immunoprecipitates (co-IP) with FLAG-tagged EPHB2 in HEK 293T cells. Based on differences in EPHB2 and EPHA3 interactomes, we reasoned that EPHA3 may serve as a negative control for the EPHB2-MYCBP2 association (Huttlin et al., 2021). We found that MYCBP2 coprecipitated with affinity-purified EPHB2-FLAG, but not EPHA3 (Fig. 1C). This confirmed MYCBP2 binding to EPHB2 and suggested MYCPB2 displays EPH receptor subtype specificity. Importantly, the EPHB2-MYCBP2 interaction was reduced by 24.5% and 53% following ephrin-B1 and ephrin-B2 treatment respectively in HeLa EPHB2 cells, suggesting the involvement of MYCBP2 in ephrin-B:EPHB2 signalling (Fig. 1D, E, eB1-Fc, p=0.1365; eB2-Fc, p=0.0002). To test this interaction in vivo, we performed co-IP using dissociated rat cortical neurons, which further confirmed MCYBP2 association with EPHB2 (Fig. 1F).

We next asked whether the kinase activity of EPHB2 is required for the formation of the EPHB2-MYCBP2 complex. To test this, we expressed GFP-tagged wild type (WT) EPHB2 or EPHB2 with a kinase-dead mutation (KD) in HeLa cells. We found that MYCBP2 showed comparable coprecipitation with WT and KD EPHB2 (Fig. 1G). Thus, EPHB2 kinase activity is not required for the formation of MYCBP2 - EPHB2 complexes, suggesting that MYCBP2 association with EPHB2 may be a prelude to ephrin-B-evoked EPHB2 signalling.

FBXO45 enhances the association between MYCBP2 and EPHB2

Our proteomics screen for EPHB2 interactors also identified FBXO45, the F-box protein that forms a ubiquitin ligase complex with MYCBP2 (Sharma et al., 2014;Saiga et al., 2009). Thus, we initially reasoned that FBXO45 might perform a similar role in the formation of the MYCBP2-EPHB2 complex. We first tested whether FBXO45 binds to EPHB2 by co-expressing MYC-tagged FBXO45 with EPHB2-FLAG or EPHA3-FLAG in HEK293T cells. Co-IP revealed that FBXO45 can associate with EPHB2 but not EPHA3, suggesting that EPHB2, FBXO45 and MYCBP2 could form a ternary complex (Fig. 2A). To test this idea, we co-expressed EPHB2-FLAG and GFP-MYCBP2 in the presence or absence of MYC-FBXO45 and examined co-IP efficiency between EPHB2 and MYCBP2. Interestingly, FBXO45 enhanced the interaction between EPHB2 and MYCBP2 (Fig. 2B, C, p=0.0068). Together, these data suggest that EPHB2 can form a complex with both MYCBP2 and FBXO45, and that FBXO45 increases the efficiency of MYCBP2-EPHB2 interaction.

Mapping binding regions for EPHB2-MYCBP2 reveals role of FBXO45 in this interaction.

A) Co-IP of EPHB2-FLAG with MYC-FBXO45 using transfected HEK293 cells. EPHB2 co-precipitates FBXO45, but EPHA3 does not. Asterisk indicates MYC-FBXO45. B) In HEK 293T cells, FBXO45 overexpression enhances EPHB2-MYCBP2 binding. Asterisk indicates MYC-FBXO45. C) Quantification of the association intensity of MYCBP2 and EPHB2 upon FBXO45 overexpression (EPHB2, p=0.0068; EPHB2 vs EPHA3, p=0.0005; one sample t-test). Error bars represent SD. D) Schematic representation of MYCBP2 N-terminal, Central and C-terminal fragments. E) Co-IP of EPHB2-FLAG with GFP-MYCBP2 fragments in HEK 293T cells. EPHB2 coprecipitates with MYCBP2 central fragment. Asterisks indicate GFP-MYCBP2 fragments. F) Schematic of chimeric domain swapping of EPHB2 (orange) and EPHA3 (grey). G) Co-IP of MYC-FBXO45 and endogenous MYCBP2 with EPHB2/EPHA3 domain swapped chimeras. H) Schematic representation of EPHB2 ΔECD (extracellular domain, aa deletions of 19-530) and ΔICD (intracellular domain, aa deletions of 590-986) truncations. I) Co-IP of endogenous MYCBP2 with EPHA3, EPHB2 and EPHB2 truncation mutants. ECD, extracellular domain; TM, transmembrane; JM, juxtamembrane; SAM, Sterile alpha motif; PBM, PDZ (PSD-95, Dlg1, Zo-1) binding motif.

Biochemical mapping of EPHB2-MYCBP2 interaction

To identify the MYCBP2 protein domain(s) required for the formation of the ternary complex with EPHB2 and FBXO45, we co-expressed EPHB2-FLAG and three GFP-MYCBP2 fragments in HEK293T cells. Co-IP revealed that the central region of MYCBP2 was sufficient for binding with EPHB2 (Fig. 2D, E). In addition, co-expression of MYC-FBXO45 demonstrated that the association of the central domain of MYCBP2 with EPHB2 is enhanced by FBOX45 and is consistent with the presence of an FBXO45 binding site within this MYCBP2 fragment (Fig. 2E) (Sharma et al., 2014).

Earlier version: To identify the domains of EPHB2 required for the formation of the tripartite complex, we took advantage of the observation that EPHA3 does not readily form a complex with MYCBP2 or FBXO45. Thus, we performed domain swapping experiments between EPHB2 and EPHA3 reasoning that placing EPHA3-specific sequences in EPHB2 would inhibit the formation of the tripartite complex. We constructed a series of FLAG-EPHB2/EPHA3 chimeras and determined whether they bound MYCBP2 in the presence of FBXO45 (Fig. 2F). Co-IP revealed that the ability of a particular chimera to associate with MYCBP2 was correlated with its association with FBXO45, in line with the MYCBP2-FBXO45 complex interacting with EPHB2. Surprisingly, we found that EPHB2-EPHA3 chimeras with an EPHA3 identity of intracellular juxtamembrane, kinase, SAM or PDZ binding domains retained the ability to associate with FBXO45 and MYCBP2, suggesting that the formation of the tripartite complex is driven by the extracellular domain and/or the transmembrane domain of EPHB2 (Fig. 2G). However, these results are also consistent with the possibility that the EPHB2 identity of the extracellular fragments could alter the conformation of the intracellular domains of EPHA3 identity, allowing the interaction with FBXO45-MYCBP2 to occur. To exclude this possibility, we created EPHB2 mutants lacking the intracellular or extracellular domains and tested their ability to complex with MYCBP2 and FBXO45 by co-IP (Fig. 2H). In these experiments only the deletion mutant lacking the intracellular domain retained its ability to form the tripartite complex. Collectively, these results argue that the combination of extracellular and transmembrane domains of EPHB2 are necessary and sufficient for formation of the MYCBP2–FBXO45–EPHB2 complex (Fig. 2I).

Revised version: To identify the domains of EPHB2 required for the formation of the tripartite complex, we took advantage of the observation that EPHA3 does not readily form a complex with MYCBP2 or FBXO45. Thus, we performed domain swapping experiments between EPHB2 and EPHA3 reasoning that placing EPHA3-specific sequences in EPHB2 would inhibit the formation of the tripartite complex. We constructed a series of FLAG-EPHB2/EPHA3 chimeras and determined whether they bound MYCBP2 in the presence of FBXO45 (Fig. 2F). Co-IP revealed that the ability of a particular chimera to associate with MYCBP2 was correlated with its association with FBXO45, in line with the MYCBP2-FBXO45 complex interacting with EPHB2. Surprisingly, we found that EPHB2-EPHA3 chimeras with an EPHA3 identity of intracellular juxtamembrane, kinase, SAM or PDZ binding domains retained the ability to associate with FBXO45 and MYCBP2, suggesting that the formation of the tripartite complex is driven by the extracellular domain and/or the transmembrane domain of EPHB2 (Fig. 2G). However, these results are also consistent with the possibility that the EPHB2 identity of the extracellular fragments could alter the conformation of the intracellular domains of EPHA3 identity, allowing the interaction with FBXO45-MYCBP2 to occur. To exclude this possibility, we created EPHB2 mutants lacking the intracellular or extracellular domains and tested their ability to complex with MYCBP2 and FBXO45 by co-IP (Fig. 2H). In these experiments only the deletion mutant lacking the intracellular domain retained its ability to form the tripartite complex. Collectively, these results argue that the combination of extracellular and transmembrane domains of EPHB2 are necessary and sufficient for formation of the MYCBP2–FBXO45–EPHB2 complex (Fig. 2I). Since EPHB2 is a transmembrane protein and MYCBP2 is localised in the cytosol. these experiments suggest that the interaction between the extracellular domain of EPHB2 and MYCBP2 might be indirect and mediated by other unknown transmembrane proteins.

MYCBP2 is required for EPHB2-mediated cellular responses

Earlier version: Given the decrease in MYCBP2–EPHB2 association evoked by ephrin-B treatment (Fig. 1D, E), we next sought to determine whether MYCBP2 fulfils a specific function in ephrin-B:EPHB2 forward signalling. Thus, we infected EPHB2-FLAG HeLa cells using lentivirus containing CRISPR sgRNA targeting the MYCBP2 exon 6, and pooled MYCBP2CRISPR cells after puromycin selection (Desbois et al., 2018). Using EPHB2-FLAG HeLa cells carrying a stably integrated empty expression vector (CTRLCRISPR) as controls, we found that MYCBP2CRISPR led to a reduction in endogenous MYCBP2 protein levels (Fig. 3A).

Revised version: Given the decrease in MYCBP2–EPHB2 association evoked by ephrin-B treatment (Fig. 1D, E), we next sought to determine whether MYCBP2 fulfils a specific function in ephrin-B:EPHB2 forward signalling. Thus, we infected EPHB2-FLAG HeLa cells using lentivirus containing CRISPR sgRNA targeting the MYCBP2 exon 6, and pooled MYCBP2CRISPR cells after puromycin selection (Fig. 3A). Using EPHB2-FLAG HeLa cells carrying a stably integrated empty expression vector (CTRLCRISPR) as controls, we found that MYCBP2CRISPR led to a reduction in endogenous MYCBP2 protein levels (Fig. 3B).

MYCBP2 CRISPR HeLa cells exhibit reduced ephrin-B2 evoked cell retraction and ephrin-B2 stripe avoidance.

Earlier version: A) MYCBP2 is reduced in HeLa MYCBP2CRISPR cells generated by sgRNA targeting MYCBP2 exon 6. B) Representative images of cell collapse assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells that were stimulated with ephrin-B2. Red arrows indicate rounded/collapsed cells. C) Quantification of collapsed cells. Statistical significance between CTRLCRISPR and MYCBP2CRISPR cells was determined using two-way ANOVA followed by Sidak’s multiple comparison test (CTRLCRISPR vs MYCBP2 CRISPR: Fc, p=0.9841; eB2-Fc, p=0.0017. Fc vs eB2-Fc: CTRLCRISPR, p<0.0001; MYCBP2 CRISPR, p=0.0115). D) Representative time-lapse sequences of CTRLCRISPR and MYCBP2CRISPR HeLa cells after ephrin-B2 treatment. E) Quantification of cell area reduction after 60 min exposure to ephrin-B2 (p=0.0268, two-tailed unpaired t test). F) Ephrin-B2 stripe assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells. Cells are visualized with Phalloidin 488 staining and nuclei are stained with DAPI (black, Fc stripes; red, ephrin-B2 stripes). G) Quantification of cells present on ephrin-B2 stripes (%). Statistical significance was determined using two-tailed unpaired t-test (p=0.0109). Error bars represent SD.

Revised version: A) Schematic of generation of stable CTRLCRISPR or MYCBP2CRISPR HeLa EPHB2-FLAG cells. Note that these are not clonal cell lines. B) MYCBP2 is reduced in HeLa MYCBP2CRISPR cells generated by sgRNA targeting MYCBP2 exon 6. C) Representative images of cell collapse assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells that were stimulated with ephrin-B2. Red arrows indicate rounded/collapsed cells. D) Quantification of collapsed cells. Statistical significance between CTRLCRISPR and MYCBP2CRISPR cells was determined using two-way ANOVA followed by Sidak’s multiple comparison test (CTRLCRISPR vs MYCBP2 CRISPR: Fc, p=0.9903; eB2-Fc, p<0.0001. Fc vs eB2-Fc: CTRLCRISPR, p<0.0001; MYCBP2 CRISPR, p=0.0003). E) Representative time-lapse sequences of CTRLCRISPR and MYCBP2CRISPR HeLa cells after ephrin-B2 treatment. F) Quantification of cell area reduction after 60 min exposure to ephrin-B2. Cell area contraction ratio: CTRLCRISPR, 27.1%; MYCBP2CRISPR, 18.8%. p=0.0268, two-tailed unpaired t test. Data points corresponding to cells in representative images in panel E are in blue. G) Ephrin-B2 stripe assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells. Cells are visualized with Phalloidin 488 staining and nuclei are stained with DAPI (black, Fc stripes; red, ephrin-B2 stripes). H) Quantification of cells present on ephrin-B2 stripes (%). Statistical significance was determined using two-tailed unpaired t-test (p=0.0109). Error bars represent SD.

Earlier version: To study the role of MYCBP2 in EphB signalling, we took advantage of an experimental paradigm in which exposure to ephrin-B2 evokes cytoskeletal contraction of HeLa cells expressing EPHB2 (Lahaie et al., 2019). Thus, following induction of EPHB2 expression, MYCBP2CRISPR and CTRLCRISPR cells were stimulated with pre-clustered Fc control or ephrin-B2 for 15 minutes and scored as collapsed or uncollapsed. While the proportion of collapsed cells for the two cell lines treated with Fc was similar (CTRLCRISPR, 5.4%; MYCBP2CRISPR, 5.7%; p=0.9841), ephrin-B2 treatment resulted in the collapse of 19.0% of CTRLCRISPR cells but only 12.1% of MYCBP2CRISPR cells (n=9 coverslips; Fig. 3B, C; p=0.0017). Moreover, when compared to the Fc conditions, MYCBP2CRISPR cells exhibited a less drastic change in collapse rate upon ephrin-B2 treatment (Fig. 3B, C; eB2 vs Fc: CTRLCRISPR, p<0.0001; MYCBP2CRISPR, p=0.0115).

Revised version: To study the role of MYCBP2 in EphB signalling, we took advantage of an experimental paradigm in which exposure to ephrin-B2 evokes cytoskeletal contraction of HeLa cells expressing EPHB2 (Lahaie et al., 2019). Thus, following induction of EPHB2 expression, MYCBP2CRISPR and CTRLCRISPR cells were stimulated with pre-clustered Fc control or ephrin-B2 for 15 minutes and scored as collapsed or uncollapsed. While the proportion of collapsed cells for the two cell lines treated with Fc was similar (CTRLCRISPR, 5.4%; MYCBP2CRISPR, 5.2%; p=0.9903), ephrin-B2 treatment resulted in the collapse of 19.9% of CTRLCRISPR cells but only 12.1% of MYCBP2CRISPR cells (n=9 coverslips; Fig. 3C, D; p<0.0001). Moreover, when compared to the Fc conditions, MYCBP2CRISPR cells exhibited a less drastic change in collapse rate upon ephrin-B2 treatment (Fig. 3C, D; eB2 vs Fc: CTRLCRISPR, p<0.0001; MYCBP2CRISPR, p=0.0003).

Earlier version: In addition, time-lapse imaging of MYCBP2CRISPR and CTRLCRISPR cells transiently transfected with an EPHB2-GFP expression plasmid revealed a similar attenuation of ephrin-B2-induced cellular contraction (Fig. 3D, E; p=0.0268). These data argue that MYCBP2 regulates a short-term cellular response evoked by ephrin-B2:EPHB2 signalling.

Revised version: In addition, time-lapse imaging of MYCBP2CRISPR and CTRLCRISPR cells transiently transfected with an EPHB2-GFP expression plasmid revealed a similar attenuation of ephrin-B2-induced cellular contraction (Fig. 3E, F; p=0.0268). These data argue that MYCBP2 regulates a short-term cellular response evoked by ephrin-B2:EPHB2 signalling.

Earlier version: To study longer-term cellular responses evoked by ephrin-B2:EPHB2 signalling, we turned to a stripe assay in which the preference of CTRLCRISPR and MYCBP2CRISPR cells for immobilized ephrin-B2 or Fc was measured. To do this, cells of either line were deposited over alternating stripes of ephrin-B2 or Fc, and stripe preference was scored for individual cells after overnight incubation. While only 33.2% of CTRLCRISPR cells resided on ephrin-B2 stripes, this proportion was significantly increased to 47.4% for MYCBP2CRISPR cells, suggesting the loss of MYCBP2 function led to a decreased repulsion from ephrin-B2 stripes (Fig. 3F, G; n=5 and 7 carpets respectively, p=0.0109). When cells were plated on Fc:Fc stripes, CTRLCRISPR and MYCBP2CRISPR cells exhibited no preference over cy3-conjugated Fc stripes (49.18% vs 51.88%, p=0.5386, images not shown). Native HeLa cells only respond to ephrin-B2 once they are made to express EphB2. Thus, our data suggest that MYCBP2 is required for EPHB2-mediated cellular responses in HeLa cells.

Revised version: To study longer-term cellular responses evoked by ephrin-B2:EPHB2 signalling, we turned to a stripe assay in which the preference of CTRLCRISPR and MYCBP2CRISPR cells for immobilized ephrin-B2 or Fc was measured. To do this, cells of either line were deposited over alternating stripes of ephrin-B2 or Fc, and stripe preference was scored for individual cells after overnight incubation. While only 33.2% of CTRLCRISPR cells resided on ephrin-B2 stripes, this proportion was significantly increased to 47.4% for MYCBP2CRISPR cells, suggesting the loss of MYCBP2 function led to a decreased repulsion from ephrin-B2 stripes (Fig. 3G, H; n=5 and 7 carpets respectively, p=0.0109). When cells were plated on Fc:Fc stripes, CTRLCRISPR and MYCBP2CRISPR cells exhibited no preference over cy3-conjugated Fc stripes (49.18% vs 51.88%, p=0.5386, images not shown). Native HeLa cells only respond to ephrin-B2 once they are made to express EphB2. Thus, our data suggest that MYCBP2 is required for EPHB2-mediated cellular responses in HeLa cells.

Loss of MYCBP2 decreases cellular levels of EPHB2 protein

The association of EPHB2 with MYCBP2 and its substrate recognition protein FBXO45 suggests that the MYCBP2 ubiquitin ligase complex could target EPHB2 for degradation, a mechanism frequently deployed to terminate transmembrane receptor signalling (Foot et al., 2017). However, the results of our cellular assays contradicted this model, and rather suggested that loss of MYCBP2 function decreased EPHB2 signalling. To further evaluate these two scenarios, we first compared EPHB2 protein levels in HeLa MYCBP2CRISPR cells and CTRLCRISPR cells by using tetracycline to induce the EPHB2 overexpression. We found that MYCBP2 loss reduced EPHB2 protein levels (Fig. 4A). To further confirm this, instead of inducing EPHB2 overexpression with tetracycline, we transfected EPHB2-FLAG plasmid into both CTRLCRISPR and MYCBP2CRISPR cells and examined EPHB2-FLAG levels after two days. As shown in Fig. 4B, C, levels of EPHB2-FLAG were significantly lowered by 26.1% in HeLa MYCBP2CRISPR cells compared to CTRLCRISPR cells (p=0.0046). We further investigated whether MYCBP2 affects EPHB2 protein turnover when cycloheximide is added to prevent new protein synthesis. EPHB2-FLAG expression was induced by tetracycline for 12 hours followed by protein synthesis inhibition with cycloheximide. We found that EPHB2 half-life was reduced in HeLa cells lacking MYCBP2 compared to control (Fig. 4D, E; 8h treatment, p=0.0474).

MYCBP2 loss-of-function increases EPHB2 protein turnover in HeLa cells.

Earlier version: A) Induced EPHB2-FLAG expression is reduced in MYCBP2CRISPR HeLa cells. B) Western blotting for transfected EPHB2-FLAG in CTRLCRISPR and MYCBP2CRISPR cells. C) Quantification of transfected EPHB2-FLAG levels (p=0.0046, one-sample t-test). D) Representative western blot of EPHB2-FLAG in CTRLCRISPR and MYCBP2CRISPR cells treated with DMSO or cycloheximide for 4h and 8h. E) Quantification of EPHB2-FLAG turnover with cycloheximide (CTRLCRISPR vs. MYCBP2CRISPR at 8h eB2-Fc stimulation, p= 0.0474, two-way ANOVA followed by Tukey’s multiple comparison test). F) Western blot showing EPHB2-FLAG degradation when cells are challenged with ephrin-B2 (1µg/ml) for different periods of time. G) Quantification of ephrin-B2-evoked EPHB2 degradation (ns, not significant; two-way ANOVA followed by Tukey’s multiple comparison test). H) Western blot of EPHB2 ubiquitination in CTRLCRISPR and MYCBP2CRISPR cells. I) Quantification of ubiquitinated EPHB2. CTRLCRISPR cells stimulated with Fc vs. eB2-Fc, p= 0.1349 (One sample t-test); MYCBP2CRISPR cells stimulated with Fc vs. EB2-Fc, p= 0.0195 (Unpaired two-tailed t-test). J) After tetracycline induction of EPHB2-FLAG expression for 16hrs, CTRLCRISPR and MYCBP2CRISPR HeLa cells were treated with DMSO (1:500) or inhibitors of the proteasome (MG132 50µM) or lysosome (BafA1 0.2µM; CoQ 50µM) for 6 hours, and EPHB2 levels were analysed by western blotting. K) Quantification of EPHB2 levels following treatment with proteasome or lysosome inhibitors. Statistical significance for the comparison between CTRLCRISPR cells treated with DMSO or inhibitors was determined by one-sample t-test (MG132, p=0.0598; BafA1, p=0.0200; CoQ, p=0.3632), whereas statistical significance for the comparison between MYCBP2CRISPR cells treated with DMSO and individual inhibitors was determined by two-tailed paired t-test (MG132, p=0.8893; BafA1, p=0.0361; CoQ, p=0.0835). L) GFP-EPHB1 and HA-EPHB3 transfected into CTRLCRISPR and MYCBP2CRISPR HeLa cells and detected by Western blot. M) Quantification of GFP-EPHB1 and HA-EPHB3 levels (EPHB1, p=0.0588; EPHB3, p=0.4253; one-sample t-test). N) FLAG-EPHA3 transfected into CTRLCRISPR and MYCBP2CRISPR HeLa cells and detected by WB. O) Quantification of FLAG-EPHA3 (p=0.5369, one-sample t-test). Error bars represent SD.

Revised version: A) Induced EPHB2-FLAG expression is reduced in MYCBP2CRISPR HeLa cells. B) Western blotting for transfected EPHB2-FLAG in CTRLCRISPR and MYCBP2CRISPR cells. C) Quantification of transfected EPHB2-FLAG levels (p=0.0046, one-sample t-test). D) Representative western blot of EPHB2-FLAG in CTRLCRISPR and MYCBP2CRISPR cells treated with DMSO or cycloheximide for 4h and 8h. E) Quantification of EPHB2-FLAG turnover with cycloheximide (CTRLCRISPR vs. MYCBP2CRISPR at 8h eB2-Fc stimulation, p= 0.0474, two-way ANOVA followed by Tukey’s multiple comparison test). F) Western blot showing EPHB2-FLAG degradation when cells are challenged with ephrin-B2 (1µg/ml) for different periods of time. G) Quantification of ephrin-B2-evoked EPHB2 degradation (ns, not significant; two-way ANOVA followed by Tukey’s multiple comparison test). Although not significant, there is an apparent trend towards lower EPHB2 levels in MYCBP2CRISPR cells, which could become significant with additional replicates. H) Western blot of EPHB2 ubiquitination in CTRLCRISPR and MYCBP2CRISPR cells. I) Quantification of ubiquitinated EPHB2. CTRLCRISPR cells stimulated with Fc vs. eB2-Fc, p= 0.1349 (One sample t-test); MYCBP2CRISPR cells stimulated with Fc vs. EB2-Fc, p= 0.0195 (Unpaired two-tailed t-test). J) After tetracycline induction of EPHB2-FLAG expression for 16hrs, CTRLCRISPR and MYCBP2CRISPR HeLa cells were treated with DMSO (1:500) or inhibitors of the proteasome (MG132 50µM) or lysosome (BafA1 0.2µM; CoQ 50µM) for 6 hours, and EPHB2 levels were analysed by western blotting. K) Quantification of EPHB2 levels following treatment with proteasome or lysosome inhibitors. Statistical significance for the comparison between CTRLCRISPR cells treated with DMSO or inhibitors was determined by one-sample t-test (MG132, p=0.0598; BafA1, p=0.0200; CoQ, p=0.3632), whereas statistical significance for the comparison between MYCBP2CRISPR cells treated with DMSO and individual inhibitors was determined by two-tailed paired t-test (MG132, p=0.8893; BafA1, p=0.0361; CoQ, p=0.0835). L) GFP-EPHB1 and HA-EPHB3 transfected into CTRLCRISPR and MYCBP2CRISPR HeLa cells and detected by Western blot. M) Quantification of GFP-EPHB1 and HA-EPHB3 levels (EPHB1, p=0.0588; EPHB3, p=0.4253; one-sample t-test). N) FLAG-EPHA3 transfected into CTRLCRISPR and MYCBP2CRISPR HeLa cells and detected by WB. O) Quantification of FLAG-EPHA3 (p=0.5369, one-sample t-test). Error bars represent SD.

Ephrin ligand treatment eventually results in Eph receptor degradation, a process associated with signalling termination. We therefore asked whether ligand-mediated EPHB2 receptor degradation depends on MYCBP2. We induced EPHB2 expression in CTRLCRISPR and MYCBP2CRISPR cells, and exposed them to ephrin-B2 for a different length of time. Western blotting revealed that 4-8 hours stimulation with ephrin-B2 reduced EPHB2 levels in CTRLCRISPR cells, but this effect was more drastic in MYCBP2CRISPR cells (Fig. 4F, G). Taken together, our results are not consistent with MYCBP2 ubiquitinating EPHB2 and causing its degradation. Unexpectedly, our data indicate that MYCBP2 stabilizes EPHB2 in HeLa cells under both naïve and ligand-challenged conditions.

Loss of MYCBP2 enhances ligand-induced EPHB2 receptor ubiquitination

Ubiquitination of receptor tyrosine kinases, including Eph receptors, can herald their degradation via the proteasome and thus termination of signalling (Haglund and Dikic, 2012;Sabet et al., 2015). This model is not supported by our results, which suggest that MYCBP2 is required for EPHB2 protein maintenance. Nonetheless, we investigated whether EPHB2 receptor ubiquitination is altered in HeLa cells depleted of MYCBP2. HeLa CTRLCRISPR and MYCBP2CRISPR cells were transfected with HA-tagged ubiquitin, EPHB2 expression was induced with tetracycline, and cells were treated with ephrin-B2 for 30min. We observed that EPHB2 receptor ubiquitination was not significantly increased in CTRLCRISPR cells after this short-term ligand treatment (Fig. 4H, I CTRLCRISPR, p=0.1349). In contrast, EPHB2 ubiquitination was significantly increased in MYCBP2CRISPR cells (Fig. 4H, I MYCBP2CRISPR, p=0.0195). This effect argues against the concept that EPHB2 is a MYCBP2 ubiquitination substrate, and suggests that in the absence of MYCBP2 degradation of the EPHB2 receptor is enhanced due to increased ubiquitination.

Earlier version: Loss of MYCBP2 promotes EPHB2 degradation through the lysosomal pathway

Revised version: A potential involvement of the lysosomal pathway in EPHB2 degradation caused by the loss of MYCBP2

Earlier version: EPHB2 can be degraded by either a proteasomal or lysosomal pathway depending on the cellular context (Cissé et al., 2011; Litterst et al., 2007;Fasen et al., 2008). Thus, we wanted to shed light on how EPHB2 is degraded and understand why EPHB2 degradation is enhanced by MYCBP2 loss of function. To do so, we induced EPHB2 expression in CTRLCRISPR and MYCBP2CRISPR cells and applied the S26 proteasome inhibitor MG132, or the lysosomal inhibitors BafilomycinA1 or Chloroquine. We found that MG132 did not have significant effects on EPHB2 levels in both cell types (Fig. 4J, K). However, we found that BafilomycinaA1 (BafA1) significantly increased EPHB2 protein levels in both HeLa CTRLCRISPR cells and MYCBP2CRISPR cells by 19% and 40%, respectively (Fig. 4J, K). We also observed a trend towards increased EPHB2 levels with Chloroquine (coQ) treatment in CTRLCRISPR (14%) and MYCBP2CRISPR (35%), further suggesting a role for lysosomal degradation (Fig. 4J, K). These observations suggest that the loss of MYCBP2 promotes EPHB2 receptor degradation through the lysosomal pathway, which expands our understanding of how lysosomal degradation of EPHB2 is regulated.

Revised version: EPHB2 can be degraded by either a proteasomal or lysosomal pathway depending on the cellular context (Cissé et al., 2011;Litterst et al., 2007;Fasen et al., 2008). Thus, we wanted to shed light on how EPHB2 is degraded and understand why EPHB2 degradation is enhanced by MYCBP2 loss of function. To do so, we induced EPHB2 expression in CTRLCRISPR and MYCBP2CRISPR cells and applied the S26 proteasome inhibitor MG132, or the lysosomal inhibitors BafilomycinA1 or Chloroquine. We found that MG132 did not have significant effects on EPHB2 levels in both cell types (Fig. 4J, K). However, we found that BafilomycinaA1 (BafA1) significantly increased EPHB2 protein levels in both HeLa CTRLCRISPR cells and MYCBP2CRISPR cells by 19% and 40%, respectively (Fig. 4J, K). We also observed a trend towards increased EPHB2 levels with Chloroquine (CoQ) treatment in CTRLCRISPR (14%) and MYCBP2CRISPR (35%), further suggesting a role for lysosomal degradation (Fig. 4J, K). Although the difference in percentage increase between CTRLCRISPR cells and MYCBP2CRISPR cells is not significant, this trend raises the possibility that the loss of MYCBP2 promotes EPHB2 receptor degradation through the lysosomal pathway.

Added in the revised version: Regulation of Eph receptor levels by MYCBP2

Removed in the revised version: Loss of MYCBP2 does not decrease EPHB1, EPHB3, or EPHA3 receptor levels

Earlier version: The above experiments raise the question of whether MYCBP2 is a general regulator of Eph receptor stability. Since EPHA3 does not form a complex with MYCBP2, and EphA receptor levels are controlled by the proteasomal pathway (Sharfe et al., 2003;Walker-Daniels et al., 2002), we hypothesized that MYCBP2 might regulate the levels of the entire EphB receptor class. To test this, we co-transfected plasmids encoding GFP-tagged EPHB1 and HA-tagged EPHB3 into HeLa CTRLCRISPR and MYCBP2CRISPR cells. Compared to CTRLCRISPR cells, EPHB1 and EPHB3 levels were not significantly reduced in MYCBP2CRISPR cells (Fig. 4L, M, EPHB1, p=0.0588; EPHB3, p=0.4253). Likewise, FLAG-EPHA3 levels were similar in CTRLCRISPR and MYCBP2CRISPR cells (Fig. 4N, O, p=0.5369). Taken together, these data suggest that MYCBP2 specifically affects EPHB2 receptor stability.

Revised version: The above experiments raise the question of whether MYCBP2 is a general regulator of Eph receptor stability. Since EPHA3 does not form a complex with MYCBP2, and EphA receptor levels are controlled by the proteasomal pathway (Sharfe et al., 2003;Walker-Daniels et al., 2002), we hypothesized that MYCBP2 might regulate the levels of the entire EphB receptor class. To test this, we co-transfected plasmids encoding GFP-tagged EPHB1 and HA-tagged EPHB3 into HeLa CTRLCRISPR and MYCBP2CRISPR cells. Compared to CTRLCRISPR cells, EPHB1 and EPHB3 levels were reduced by 57.1% and 12.5% respectively in MYCBP2CRISPR cells (Fig. 4L, M, EPHB1, p=0.0588; EPHB3, p=0.4253). Although not statistically significantly, there is an apparent trend towards a decrease in EPHB1 levels. On the other hand, FLAG-EPHA3 levels were similar in CTRLCRISPR and MYCBP2CRISPR cells (Fig. 4N, O, p=0.5369). Taken together, these data suggest that MYCBP2 may stabilize other EphB subclass receptors.

Loss of MYCBP2 attenuates the magnitude of EPHB2 cellular signalling

EPHB2 receptor activation evokes signal transduction events such as tyrosine phosphorylation of EPHB2, activation of EPHB2 tyrosine kinase function and phosphorylation of the ERK1/2 downstream effector (Poliakov et al., 2008). We therefore measured EPHB2 and ERK1/2 phosphorylation in CTRLCRISPR and MYCBP2CRISPR cells for up to 8 hours after ephrin-B2 application (Fig. 5A). P-EPHB2 (pY20) and p-ERK1/2 signals were normalised to GAPDH and ERK1/2, respectively. In HeLa CTRLCRISPR cells, pTyr-EPHB2 response reached a plateau after 1-2 h treatment and remained up 8 hours post-stimulation (Fig. 5A-C). On the contrary, ephrin-B2-evoked phosphorylation of EPHB2 was reduced in MYCBP2CRISPR cells with quantitative results showing a significant reduction by 8 hours of treatment (Fig. 5A, C; 8h p=0.0331). We were also able to detect significantly lower p-ERK1/2 levels in MYCBP2CRISPR cells relative to CTRLCRISPR cells, although activation of ERK1/2 by ephrin-B2 was variable (Fig. 5D. 4h, p=0.0494; 8h, p=0.0078; n=6). We again noted significantly enhanced EPHB2 degradation in MYCBP2CRISPR cells (Fig. 5E; 8h, p=0.0437). This decrease was also observed when CTRLCRISPR and MYCBP2CRISPR cells were treated with ephrin-B1 (data not shown).

MYCBP2 depletion impairs EPHB2 phosphorylation and ERK1/2 activation in HeLa cells.

A) Representative western blot for pERK1/2 and pTyr-EPHB2 detected in CTRLCRISPR and MYCBP2CRISPR cells treated with ephrin-B2 (eB2-Fc) for different periods (n=6). Membranes were striped and reblotted with anti-ERK1/2, anti-EPHB2, anti-GAPDH and anti-MYCBP2 antibodies as controls. B) Quantification of EPHB2 tyrosine phosphorylation in CTRLCRISPR cells evoked by ephrin-B2 treatment (15 min, p=0.1363; 30 min, p=0.2056; 1h, p=0.0342; 2h, p=0.0234; 4h, p=0.0068; 8h, p= 0.0231; one-sample t-test). C) Quantification of EPHB2 tyrosine phosphorylation in CTRLCRISPR and MYCBP2CRISPR HeLa cells (unstimulated, p=0.5589, one-sample t-test; stimulated for 15 min, p=0.7463; 30 min, p=0.5520; 1h, p=0.1920; 2h, p=0.2009; 4h, p=0.1550; 8h, p=0.0331; two-tailed unpaired t-test). D) Quantification of pERK1/2 in CTRLCRISPR and MYCBP2CRISPR HeLa cells (0 min, p=0.0168, one-sample t-test; 15 min, p=0.6695; 30 min, p= 0.6649; 1h, p= 0.1776; 2h, p= 0.1479; 4h, p= 0.0494; 8h, p= 0.0078; two-tailed unpaired t-test). E) Quantification of EPHB2 in CTRLCRISPR and MYCBP2CRISPR HeLa (0 min, p=0.4604, one-sample t-test; 15 min, p=0.2222; 30 min, p=0.1376; 1h, p=0.0651; 2h, p=0.0736; 4h, p=0.2451; 8h, p=0.0437, two-tailed unpaired t-test). F) Quantification of ephrin-B2-evoked EPHB2 tyrosine phosphorylation levels relative to total EPHB2 protein levels (0 min, p=0.7058, one-sample t-test; 15 min, p=0.2464; 30 min, p=0.7835; 1h, p=0.9164; 2h, p=0.7196; 4h, p=0.8625; 8h, p=0.5750, two-tailed unpaired t-test). G) Quantification of ephrin-B2-evoked pERK1/2 relative to EPHB2 total protein levels (0 min, p=0.3308, one-sample t-test; 15 min, p=0.3856; 30 min, p=0.0624; 1h, p=0.2683; 2h, p=0.1284; 4h, p=0.7998; 8h, p=0.7790, two-tailed unpaired t-test. Error bars represent SD.

Earlier version: Next, we asked whether the decrease in ligand-evoked EPHB1/2 and ERK1/2 activation in MYCBP2CRISPR cells reflects a requirement for MYCBP2 in the EPHB2 signalling cascade per se, or whether it is explained by decreased EPHB2 protein levels caused by MYCBP2 loss. We thus normalized p-EPHB2 and p-ERK1/2 signal to the levels of EPHB2 protein at all time points, which revealed that kinetics and magnitude of EPHB2 and ERK1/2 phosphorylation are similar between CTRLCRISPR and MYCBP2CRISPR cells (Fig. 5F, G). Although these data argue against a direct role of MYCBP2 in the early events of EPHB2 signalling, they nevertheless indicate that MYCBP2 loss results in marked attenuation of EPHB2 activation and its downstream pERK1/2 signalling in line with decreased cellular responses to ephrin-B2.

Revised version: Next, we asked whether the decrease in ligand-evoked EPHB1/2 and ERK1/2 activation in MYCBP2CRISPR cells reflects a requirement for MYCBP2 in the EPHB2 signalling cascade per se, or whether it is explained by decreased EPHB2 protein levels caused by MYCBP2 loss. We thus normalised p-EPHB2 and p-ERK1/2 signal to the levels of EPHB2 protein at all time points, which revealed that kinetics and magnitude of EPHB2 and ERK1/2 phosphorylation are similar between CTRLCRISPR and MYCBP2CRISPR cells (Fig. 5F, G). Although these data argue against a direct role of MYCBP2 in the early events of EPHB2 signalling, they nevertheless indicate that MYCBP2 loss results in marked attenuation of EPHB2 activation and its downstream pERK1/2 signalling in line with decreased cellular responses to ephrin-B2.

Exogenous Fbxo45 binding domain of MYCBP2 disrupts the EPHB2-MYCBP2 interaction

Previous studies showed that FBXO45 binds to the FBD1 domain of MYCBP2, and exogenous FBD1 overexpression can disrupt the FBXO45-MYCBP2 association (Sharma et al., 2014). We thus tested whether FBD1 overexpression can interfere with the formation of the EPHB2-MYCBP2 complex (Fig. 6A). Indeed, the expression of GFP-FBD1 wild-type (WT) in HEK cells expressing EPHB2-FLAG and MYC-FBXO45 reduced binding between EPHB2 and MYCBP2 (Fig. 6B). This effect was not observed with a GFP-FBD1 mutant (mut) fragment that harbours three point mutations that inhibit binding to FBXO45. Reduction of the EPHB2-MYCBP2 interaction was also observed in cells that were not overexpressing FBXO45 (Fig. 6C). Thus, exogenous FBD1 can specifically disrupt the EPHB2-MYCBP2 association. This suggests that FBXO45 binding to MYCBP2 may facilitate for formation of the MYCBP2-EPHB2 complex.

Exogenous FBD1 fragment of MYCBP2 disrupts EPHB2-MYCBP2 binding and impairs EPHB2 function in HeLa cells.

A) Schematic illustrating competition of exogenous MYCBP2-FBD1 fragment that disrupts MYCBP2-FBXO45 binding and leads to MYCBP2 reduction in EPHB2 complexes. B) Exogenous FBD1 WT overexpression leads to reduced EPHB2-MYCBP2 binding in HEK293 cells despite co-expression of FBXO45. C) FBD1 overexpression also disrupts EPHB2-MYCBP2 binding in the absence of FBXO45 overexpression. D) Representative images of ephrin-B2 stripe assays using HeLa cells expressing GFP-FBD1 mut or GFP-FBD1 WT. E) Quantification of cells present on eB2 stripes (p=0.0107, two-tailed unpaired t-test). Error bars represent SD.

FBD1 expression impairs EPHB2-mediated neuronal responses

To determine whether the EPHB2-MYCBP2 interaction is required for EPHB2 function, we disrupted binding via FBD1 ectopic expression in cell lines or primary neurons, and studied cellular responses to ephrin-B treatment. We introduced GFP-FBD1 mut or GFP-FBD1 WT into HeLa cells with inducible EPHB2 expression and cultured these cells on ephrin-B2 and Fc stripes. Following overnight culture, only 19.2% of cells expressing GFP-FBD1 mut resided on ephrin-B2 stripes. In contrast, 28.6% of cells expressing GFP-FBD1 WT were found on ephrin-B2 stripes indicating that FBD1 expression can dampen ephrin-B2:EPHB2 mediated cell repulsion (Fig. 6D, E; p=0.0107). In contrast, cells expressing GFP-FBD1 mut or GFP-FBD1 WT displayed no preference for either one of Fc:Fc control stripes (49.63% vs 49.85%, p=0.9560, images not shown).

Since EPHB2 is expressed in the embryonic chicken spinal cord, we performed in ovo electroporation to introduce plasmids encoding GFP-FBD1 mut or GFP-FBD1 WT into the spinal cords of 2.5-day old embryonic chickens (Hamburger-Hamilton; (HH) Stage 15-17; (Hamburger and Hamilton, 1951;Luria et al., 2008)). Two days later spinal cords were dissected, divided into explants, cultured overnight on alternating stripes containing ephrin-B2 or Fc, and axonal GFP signal present over ephrin-B2 versus Fc stripes was determined. When explants expressed GFP-FBD1 mut (negative control), we observed 26.5% of axonal GFP signal resided on ephrin-B2 stripes (Fig. 7A, B). In contrast, we found a significant increase of 32.2% of neurites expressing FBD1 WT on ephrin-B2 stripes (Fig. 7A, B; p=0.0410). Thus, FBD1 expression impairs long-term repulsive responses to ephrinB2-EPHB2 signalling in spinal explants.

Exogenous FBD1 overexpression impairs EPH receptor functions in chick spinal cord explants and mouse hippocampal neurons.

Earlier version: A) Representative images of ephrin-B2 stripe assays with chick embryonic spinal cord explants overexpressing GFP-FBD1 mut (negative control) or GFP-FBD1 WT. Images with inverted GFP signal in dark pixels on Fc (white) / eB2 (pink) stripes are placed beside the original images. B) Quantification of GFP-positive branches present on ephrin-B2 stripes (GFP-FBD1 mut vs. GFP-FBD1 WT, p=0.0410, two-tailed unpaired t-test). C) Representative images of DIV2 mouse hippocampal neurons overexpressing GFP-FBD1 mut or GFP-FBD1 WT and challenged with Fc control, ephrin-B1 (eB1-Fc) or ephrin-B2 (eB2-Fc). D) Quantification of growth cone collapse rate for hippocampal neurites. GFP-FBD1 mut: Fc vs. eB1-Fc, p=<0.0001; Fc vs. eB2-Fc, p=0.0008; GFP-FBD1 WT: Fc vs. eB1-Fc p=0.4357; Fc vs eB2-Fc p=0.0065. Two-way ANOVA followed by Tukey’s multiple comparisons test. Error bars represent SD.

Revised version: A) Representative images of ephrin-B2 stripe assays with chick embryonic spinal cord explants overexpressing GFP-FBD1 mut (negative control) or GFP-FBD1 WT. Images with inverted GFP signal in dark pixels on Fc / eB2 (pink) stripes are placed beside the original images. B) Quantification of GFP-positive neurites present on ephrin-B2 stripes (GFP-FBD1 mut vs. GFP-FBD1 WT, p=0.0410, two-tailed unpaired t-test). C) Representative images of DIV2 mouse hippocampal neurons overexpressing GFP-FBD1 mut or GFP-FBD1 WT and challenged with Fc control or ephrin-B1 (eB1-Fc). D) Quantification of growth cone collapse rate for hippocampal neurites. GFP-FBD1 mut: Fc vs. eB1-Fc, p=0.0006; GFP-FBD1 WT: Fc vs. eB1-Fc p=0.1341. Two-way ANOVA followed by Sidak’s multiple comparisons test. Error bars represent SD.

Earlier version: To study short-term neuronal responses to ephrin-Bs, we turned to mouse hippocampal neurons and a growth cone collapse assay (Srivastava et al., 2013). Here, we electroporated GFP-FBD1 mut or GFP-FBD1 WT expression plasmids into dissociated hippocampal neurons and treated them with pre-clustered Fc, ephrin-B1 or ephrin-B2 for one hour (Fig. 7C). Neurons expressing GFP-FBD1 mut showed 33.9% growth cone collapse with Fc treatment, while Ephrin-B1 and Ephrin-B2 treatment elicited significant increases to 57.9% and 53.3% collapse, respectively (Fig. 7C, D). Importantly, both ephrin-B ligands significantly increased the frequency of growth cone collapse with respect to Fc treatment when FBD1 mut was expressed (Fig. 7D; p<0.0001, and p=0.0008). In contrast, this effect was either significantly reduced or less evident in neurons expressing GFP-FBD1 WT: with only 48.1% and 57.5% growth cones being collapsed by ephrin-B1 and ephrin-B2, respectively, compared to 41.7% being collapsed by Fc treatment (Fig. 7D; p=0.4357 and p=0.0065, respectively). Together, these data indicate that impairing MYCBP2 function via FBD1 expression disrupts ephrin:B-EPHB2 signalling in axonal guidance.

Revised version: To study short-term neuronal responses to ephrin-Bs, we turned to mouse hippocampal neurons and a growth cone collapse assay (Srivastava et al., 2013). Here, we electroporated GFP-FBD1 mut or GFP-FBD1 WT expression plasmids into dissociated hippocampal neurons and treated them with pre-clustered Fc or ephrin-B1 for one hour (Fig. 7C). Neurons expressing GFP-FBD1 mut showed 22.7% growth cone collapse with Fc treatment, while ephrin-B1 treatment elicited significant increases to 35.9% collapse (Fig. 7C, D, p=0.0006). In contrast, ephrin-B1 failed to significantly induce growth cone collapse in neurons expressing GFP-FBD1 WT: with only 31% growth cones being collapsed by ephrin-B1, compared to 25.3% being collapsed by Fc treatment (Fig. 7D; p=0.1341). Together, these data indicate that impairing MYCBP2 function via FBD1 expression disrupts ephrin:B-EPHB2 signalling in axonal guidance.

Added in the revised version: Genetic interactions between the C. elegans Eph receptor and the MYCBP2 signalling network

Removed in the revised version: C. elegans ephrin receptor VAB-1 interacts genetically with known RPM-1/MYCBP2 binding proteins

Earlier version: Next, we sought to test genetic interactions between an ephrin receptor and the MYCBP2 signalling network using an in vivo animal model. To do so, we turned to C. elegans which has a sole Eph family receptor called VAB-1(George et al., 1998), and a single MYCBP2 ortholog called RPM-1(Grill et al., 2016). Previous studies have shown that RPM-1/MYCBP2 is required to terminate axon outgrowth with failed termination defects in rpm-1 mutants arising due to impaired growth cone collapse(Schaefer et al., 2000;Borgen et al., 2017). Prior studies have also demonstrated that RPM-1 functions via numerous downstream pathways, and that RPM-1 binding proteins display genetic enhancer interactions (Grill et al., 2007;Grill et al., 2012;Tulgren et al., 2014;Baker et al., 2014). Given our cell-based biochemical results showing that MYCBP2 binds EphB2, we opted to evaluate two known RPM-1 binding proteins for genetic interactions with the VAB-1 Ephrin receptor. 1) The Rab GEF GLO-4 is an RPM-1 binding protein that is orthologous to mammalian SERGEF. GLO-4 functions via the GLO-1/RAB32 small GTPase and is not involved in RPM-1 ubiquitin ligase activity (Grill et al., 2007;Sharma et al., 2014)(Fig 8A). 2) We also tested the C. elegans FBXO45 ortholog, FSN-1, which is the F-box protein of the RPM-1 ubiquitin ligase complex (Liao et al., 2004;Sharma et al., 2014;Baker and Grill, 2017;Desbois et al., 2022) (Fig 8A).

Revised version: Next, we sought to test genetic interactions between an ephrin receptor and the MYCBP2 signalling network using an in vivo animal model. To do so, we turned to C. elegans which has a sole Eph family receptor (EPHR) called VAB-1 and a single MYCBP2 ortholog called RPM-1 (Grill et al., 2016;George et al., 1998). Previous studies have shown that RPM-1/MYCBP2 is required to terminate axon outgrowth (Borgen et al., 2017;Schaefer et al., 2000). Furthermore, RPM-1 functions as a hub upstream of a number of signalling pathways (Grill et al., 2007;Grill et al., 2012;Tulgren et al., 2014;Baker et al., 2014). Genetic results from these studies demonstrated that mutants for RPM-1 binding proteins display genetic enhancer interactions with one another, but do not enhance defects when combined with rpm-1 mutants.

Added in the revised version: Given the biochemical interactions between MYCBP2, FBXO45 and EPHB2, we first evaluated genetic interactions between VAB-1/EPHR and two known RPM-1 binding proteins:

Added in the revised version: 1) Rab GEF GLO-4, an orthologue of mammalian SERGEF that functions via the GLO-1/RAB32 small GTPase and is not involved in RPM-1 ubiquitin ligase activity (Grill et al., 2007). 2) FSN-1, an orthologue of FBXO45, that is the F-box substrate selector protein of the RPM-1 ubiquitin ligase complex (Fig 8A) (Liao et al., 2004). RPM-1 genetic interactions were studied in the left and right PLM mechanosensory neurons of C. elegans, both of which extend an axon anteriorly until it terminates posterior to the cell body of the respective ALM mechanosensory neuron (Fig. 8B). This process is visualized using a transgene, muIs32 (Pmec-7::GFP), that expresses GFP in the PLM and ALM mechanosensory neurons. As observed previously, a null allele of vab-1 that deletes exons 1-4 showed a significant increase in incidence of PLM axon extension beyond the ALM cell body (overextension) compared to wild type controls (Fig. 8B, C) (Mohamed and Chin-Sang, 2006). To study the interaction between VAB-1/EPHR and GLO-4/SERGEF, we first evaluated glo-4 mutants, in which we observed two kinds of overextension defects: one where PLM axons extend past the ALM cell body in a straight line, and a more severe defect where PLM axons “hook” ventrally (Fig. 8B, C). The frequency of both types of axon termination defects were significantly enhanced in vab-1; glo-4 double mutants (Fig. 8B, C). Overextension defects were significantly rescued by transgenic expression of VAB-1 in vab-1; glo-4 double mutants (Fig. 8C). Similarly, we observed an enhanced frequency of both overextension and hook defects in vab-1; fsn-1 double mutants compared to either mutant alone (Fig. 8D). Thus, VAB-1/EPHR interacts genetically with two proteins known to bind and function downstream of RPM-1/MYCB2.

Earlier version: C. elegans VAB-1 ephrin receptor interacts genetically with two known RPM-1/MYCBP2 binding proteins FSN-1/FBXO45 and GLO-4/SERGEF.

Revised version: C. elegans VAB-1 ephrin receptor interacts genetically with known RPM-1/MYCBP2 binding proteins FSN-1/FBXO45 and GLO-4/SERGEF.

Earlier version: A) Schematic shows the known RPM-1/MYCBP2 binding proteins GLO-4/SERGEF and FSN-1/FBXO45. GLO-4 functions independent of RPM-1 ubiquitin ligase. FSN-1 is the F-box protein that forms a ubiquitin ligase complex with RPM-1. Adapted from Grill et al., 2016. B) Schematic of axon morphology and axon termination site for PLM mechanosensory neurons. Shown are representative images of failed axon termination defects observed for PLM neurons of indicated genotypes. Examples of moderate overextension defects (arrowhead) observed in vab-1/EphR and glo-4/SERGEF single mutants. Example of severe hook defects (arrow) observed in vab-1; glo-4 double mutants. C) Quantitation of axon termination defects for indicated genotypes. vab-1; glo-4 double mutants show enhanced frequency of both severe hook (black) and moderate overextension (grey) failed termination defects. Overextension defects are significantly reduced by transgenic expression of VAB-1. D) Quantitation of axon termination defects for indicated genotypes. vab-1; fsn-1 double mutants show enhanced termination defects. E) Quantitation of axon termination defect for indicated genotypes. Axon termination defects are not suppressed in vab-1; rpm-1 double mutants compared to rpm-1 single mutants. n is defined as a single count of 20-30 animals. Means are shown from 8 to 10 counts (20-30 animals per count) for each genotype, and error bars represent SEM. Significance determined using Student’s t-test with Bonferroni correction. ** p < 0.01; *** p < 0.001; n.s, not significant. Scale bar is 20 μm.

Revised version: A) A schematic showing the known RPM-1/MYCBP2 binding proteins GLO-4/SERGEF and FSN-1/FBXO45. GLO-4 functions independent of RPM-1 ubiquitin ligase. FSN-1 is the F-box protein that forms a ubiquitin ligase complex with RPM-1. Adapted from Grill et al., 2016. B) Schematic representation of axon morphology and axon termination site for PLM mechanosensory neurons and representative images of failed axon termination defects observed in PLM neurons for indicated genotypes. Axon termination visualized using muIs32 (Pmec-7::GFP), which expresses GFP in the PLM and ALM mechanosensory neurons. Examples of moderate severity overextension defects (arrowhead) observed in vab-1/EphR and glo-4/SERGEF single mutants. Example of severe overextension (hook) defects (arrow) observed in vab-1; glo-4 double mutants. C) Quantitation of axon termination defects for indicated genotypes using muIs32. vab-1; glo-4 double mutants show enhanced frequency of both hook (black) and overextension (grey) failed termination defects. Overextension defects are significantly reduced by transgenic expression of VAB-1. D) Quantitation of axon termination defects for indicated genotypes. vab-1; fsn-1 double mutants show enhanced termination defects. E) Quantitation of axon termination defects for indicated genotypes using muIs32. Axon termination defects are not suppressed in vab-1; rpm-1 double mutants compared to rpm-1 single mutants. F) zdIs5 (Pmec-4::GFP) was used to quantify axon termination defects for indicated genotypes. vab-1; fsn-1 double mutants show enhanced frequency of overextension defects (grey). Frequency and severity of axon termination defects is not significantly different between vab-1; rpm-1 double mutants and rpm-1 single mutants. n is defined as a single count of 20-30 animals. Means are shown from 8 to 10 counts (20-30 animals per count) for each genotype, and error bars represent SEM. Significance determined using Student’s t-test with Bonferroni correction for multiple comparisons. ** p < 0.01; *** p < 0.001; n.s, not significant. Scale bar is 20 μm.

Added in the revised version: Next, we examined genetic interactions between vab-1 and an rpm-1 protein null allele (AlAbdi et al., 2023). We compared the incidence of PLM overextension in vab-1; rpm-1 double mutants to rpm-1 single mutants and did not observe a significant difference between them (Fig. 8E). Because rpm-1 mutants displayed a high frequency of PLM hook defects when the PLM axon was visualized using the muIs32 transgene, we pivoted to address phenotypic saturation. To do so, we evaluated vab-1 genetic interactions using the zdIs5 (Pmec4::GFP) transgene to label PLM axons. This axonal reporter was used previously to demonstrate that the frequency of rpm-1 hook defects can be enhanced by mutations in genes that are not RPM-1/MYCBP2-binding proteins (Borgen et al., 2017). In contrast, the incidence of rpm-1 hook defects was not increased by mutations that impair RPM-1 binding proteins. Like prior findings, rpm-1 mutants on the zdIs5 background result in a lower frequency of hook defects than rpm-1 mutants on muIs32 (Fig 8F). Using the zdIs5 background, we found that vab-1; fsn-1 double mutants display a higher frequency of overextension defects when compared to either single mutant (Fig 8F). This outcome is similar to what we observed in the muls32 background (Fig 8D). Finally, when comparing the incidence of PLM hook defects in vab-1; rpm-1 double mutants to single mutants in the zdIs5 background, we did not observe any significant differences (Fig 8D).

Added in the revised version: Collectively, these results suggest two general conclusions: 1) The enhanced incidence of axon termination defects in vab-1; glo-4 and vab-1; fsn-1 double mutants compared to single mutants indicates that VAB-1/EPHR functions in parallel to known RPM-1 binding proteins to facilitate axon termination. 2) vab-1; rpm-1 double mutants do not display an increased frequency or severity of axon termination defects compared to rpm-1 single mutants on multiple transgenic backgrounds. Thus, because we are using null alleles, we conclude that VAB-1/EPHR functions in the same genetic pathway as RPM-1/MYCBP2.

Removed in the revised version: To evaluate genetic interactions between the VAB-1 Eph receptor and RPM-1/MYCBP2 binding proteins, we used the PLM mechanosensory neurons of C. elegans. There are two PLM neurons with one on either side of the body. Each PLM neuron extends a primary axon that grows anteriorly and terminates at a precise location before the cell body of another mechanosensory neuron called ALM (Fig. 8B). Consistent with previous studies (Mohamed and Chin-Sang, 2006), vab-1/EphR loss-of-function (lf) mutants presented moderate failed termination defects where the PLM axon overextends beyond the ALM cell body (Fig. 8B).

Removed in the revised version: Quantitation indicated that this phenotype occurred at a low, but significant frequency (Fig. 8C). We then examined the genetic relationship between vab-1 and the RPM-1 binding protein glo-4/SERGEF. In vab-1; glo-4 double mutants, we observed both moderate overextension defects and more severe failed termination defects where the axon hooked ventrally (Fig. 8B). The frequency of both types of axon termination defects were significantly enhanced in vab-1; glo-4 double mutants (Fig. 8B, C). Overextension defects were significantly rescued by transgenic expression of VAB-1 in vab-1; glo-4 double mutants (Fig. 8C). We also tested the relationship between the F-box protein fsn-1/FBXO45 and vab-1, and observed enhanced defects in vab-1; fsn-1 double mutants (Fig. 8D). Finally, we evaluated vab-1; rpm-1 double mutants and did not observe a significant change in phenotypic severity or penetrance compared to rpm-1 single mutants (Fig. 8E).

Removed in the revised version: Collectively, these genetic experiments support two conclusions. 1) Enhanced defects in vab-1; glo-4 and vab-1; fsn-1 double mutants indicate that the VAB-1 Ephrin receptor functions in parallel to known RPM-1 binding proteins to facilitate axon termination. These results are consistent with the VAB-1 Ephrin receptor functioning within the RPM-1 signalling network. 2) Lack of suppression in vab-1; fsn-1 and vab-1; rpm-1 double mutants indicates that VAB-1 is not inhibited or degraded by the RPM-1/FSN-1 ubiquitin ligase complex. Thus, our findings in C. elegans are consistent with our biochemical studies in mammalian cells that demonstrate MYCBP2 stabilizes EPHB2 rather than degrading it. We note that because C. elegans contains only a single Ephrin receptor, our findings with C. elegans do not necessarily pertain specifically to EPHB2. Nonetheless, these results add further evidence indicating there are substantial genetic interactions between known MYCBP2 binding proteins and Eph receptors.

Discussion

Our MS-based proteomics efforts to identify EPHB2 interacting proteins yielded MYCBP2, a signalling hub and ubiquitin ligase that is functionally linked to many of the cellular processes also mediated by EPHB2, including cellular growth, proliferation, synapse formation, and axon development. Our experiments argue against EPHB2 being a MYCBP2 ubiquitination substrate. Instead, we envisage a model where MYCBP2 controls EPHB2 signalling indirectly by preventing its lysosomal degradation and maintaining EPHB2 protein levels sufficiently high to mediate efficient cellular and axonal growth cone repulsion from ephrin-B ligands. The interaction between MYCBP2 and EPHB2 may allow the coupling EPHB2 to fundamental cellular processes that control growth, proliferation, and survival. Here, we discuss the molecular logic of MYCBP2 and EPHB2 association in the context of Eph receptor signalling, its potential implications for neural development and diversification of EphB2 signalling.

Functional significance of MYCBP2-EPHB2 complex formation

Earlier version: Our biochemical experiments validate the formation of a MYCBP2-EPHB2 complex and suggest that this association is decreased following ligand application. Because PHR proteins like MYCBP2 are large signalling hubs, MYCBP2 association with EPHB2 might sterically hinder the formation of EPHB2 multimers and clusters necessary for signalling. Thus, ligand-induced dissociation of the MYCBP2-EPHB2 complex could be a prelude to signalling. Our findings indicate that MYCBP2 association with EPHB2 is enhanced by FBXO45, a subunit of the MYCBP2/FBXO45 complex that mediates ubiquitination substrate binding. Previous work showed that MYCBP2 functions to polyubiquitinate specific protein substrates, targeting them for degradation and inhibition (Crawley et al., 2019;Han et al., 2012;Nakata et al., 2005;Desbois et al., 2022). We considered the possibility that the ubiquitin ligase activity of MYCBP2 is important for the termination of EPHB2 signalling, but several lines of evidence argue against this idea: (1) the application of ephrin-B2 ligand results in the dissociation of the MYCBP2-EPHB2 complex (Walker-Daniels et al.), (2) loss of MYCBP2 function results in decreased cellular levels of EPHB2 protein, and (3) impairing MYCBP2 increases ligand-stimulated EPHB2 ubiquitination. In addition, our quantification of EPHB2 signalling suggested that decreased EPHB2 and ERK phosphorylation seen in MYCBP2-deficient cells can be accounted for by the decrease in EPHB2 receptor levels. Thus, a more plausible model that is consistent with our results is that MYCBP2 association with EPHB2 protects EPHB2 from turnover by lysosome-mediated degradation and prevents EPHB2 ubiquitination by the action of unidentified ubiquitin ligases. Interestingly, studies in C. elegans have shown that RPM-1/MYCBP2 regulates lysosome biogenesis via the GLO-4/GLO-1 pathway (Grill et al., 2007). While this established a link between MYCBP2 signalling and the endo-lysosomal degradation system, our results now indicate that MYCBP2 can also influence turnover of EPHB2 via lysosomal degradation. Consistent with these findings, we observed that the Eph receptor VAB-1 functions in parallel genetic pathways with the known RPM-1 binding protein GLO-4/SERGEF that regulates lysosome biogenesis.

Revised version: Our biochemical experiments validate the formation of a MYCBP2-EPHB2 complex and suggest that this association is decreased following ligand application. Because PHR proteins like MYCBP2 are large signalling hubs, MYCBP2 association with EPHB2 might sterically hinder the formation of EPHB2 multimers and clusters necessary for signalling. Thus, ligand-induced dissociation of the MYCBP2-EPHB2 complex could be a prelude to signalling. Our findings indicate that MYCBP2 association with EPHB2 is enhanced by FBXO45, a subunit of the MYCBP2/FBXO45 complex that mediates ubiquitination substrate binding. Previous work showed that MYCBP2 functions to polyubiquitinate specific protein substrates, targeting them for degradation and inhibition (Crawley et al., 2019;Han et al., 2012;Nakata et al., 2005;Desbois et al., 2022). We considered the possibility that the ubiquitin ligase activity of MYCBP2 is important for the termination of EPHB2 signalling, but several lines of evidence argue against this idea: 1) the application of ephrin-B2 ligand results in the dissociation of the MYCBP2-EPHB2 complex, 2) loss of MYCBP2 function results in decreased cellular levels of EPHB2 protein, and 3) impairing MYCBP2 increases ligand-stimulated EPHB2 ubiquitination. In addition, our quantification of EPHB2 signalling suggested that decreased EPHB2 and ERK phosphorylation seen in MYCBP2-deficient cells can be accounted for by the decrease in EPHB2 receptor levels. Thus, a more plausible model that is consistent with our results is that MYCBP2 association with EPHB2 protects EPHB2 from turnover by lysosome-mediated degradation and prevents EPHB2 ubiquitination by the action of unidentified ubiquitin ligases. Interestingly, studies in C. elegans have shown that RPM-1/MYCBP2 regulates lysosome biogenesis via the GLO-4/GLO-1 pathway (Grill et al., 2007). While this established a link between MYCBP2 signalling and the endo-lysosomal degradation system, our results now indicate that MYCBP2 can also influence turnover of EPHB2 via lysosomal degradation. Our C. elegans experiments reveal complex genetic interactions between VAB-1/EPHR and members of the RPM-1/MYCBP2 signalling network. These are consistent with VAB-1/EPHR and RPM-1/MYCBP2 acting in the same pathway, while suggesting that VAB-1/EPHR acts in a parallel genetic pathway with GLO-4, the RPM-1 binding protein that mediates effects on lysosome biogenesis. As a caveat, we note that because C. elegans contains only a single Eph receptor, our findings with do not necessarily pertain specifically to EPHB2.

The “protective” effect of MYCBP2 vis-à-vis EPHB2 ubiquitination and lysosomal degradation might be secondary to effects on other aspects of EPHB2 signalling that we have not explored experimentally. For example, since MYCBP2 is a signalling hub with multiple substrates and binding proteins, it could bring EphB2 into close proximity to components of the MYCBP2 signalling network integrating cell-cell communication via ephrin:Eph signals with MYCBP2 intracellular signalling to influence fundamental cellular processes. One example of this is suggested by a recent study demonstrating a link between EPHB2 and mTOR-mediated cell growth signalling pathways (Gagne et al., 2021). One potential mechanism could involve Ephrin binding-induced dissociation of the EPHB2-MYCBP2 interaction, allowing MYCBP2 ubiquitination of the TSC complex that regulates mTOR function (Han et al., 2012).

The formation of a MYCBP2-EPHB2 complex that includes other signalling receptors could also explain the curious result that the extracellular domain of EPHB2 is critical for MYCBP2 association. The finding would be consistent with the extracellular domain of EPHB2 interacting with the extracellular domains of other receptors, whose intracellular domains are linked more directly with MYCBP2 and FBXO45. An alternative explanation may be a non-classical extracellular MYCBP2-EPHB2 interaction similar to the one proposed for the extracellular domain of N-Cadherin and FBXO45 (Na et al., 2020).

MYCBP2 and EPHB2 functions in the developing nervous system

Our proteomic, biochemical, and genetic experiments indicate that MYCBP2 and EPHB2 function in the same pathway. This is also supported by the striking similarity of MYCBP2 and EPHB2 loss of function phenotypes in the mouse nervous system. In the context of developing neuronal connections, EphB2 and Mycbp2/Phr1 mouse mutants display similar axon guidance phenotypes such as abnormal limb nerve trajectories by motor axons, defective growth cone crossing of the midline at the level of the optic chiasm and decreased connectivity between the two cortical hemispheres through the corpus callosum (Lewcock et al., 2007;Luria et al., 2008;Williams et al., 2003;Henkemeyer et al., 1996;D’Souza et al., 2005). Both proteins also function in synaptic development and their loss of function leads to decreased numbers of synapses and altered synapse morphology (Wan et al., 2000;Bloom et al., 2007;Dalva et al., 2000;Zhen et al., 2000). Furthermore, outside the nervous system, both are involved in a variety of cancers with recent evidence linking them to esophageal adenocarcinoma and c-MYC-dependent control of cell proliferation (Venkitachalam et al., 2022;Han et al., 2012;Genander et al., 2009). Our proteomic, biochemical, and cellular experiments together with genetic interaction studies in C. elegans now provide a new framework in which to consider phenotypic and disease links between EPHB2 and MYCBP2.

Earlier version: Importantly, the biomedical relevance of our findings are heightened by a recent study that identified genetic variants in MYCBP2, which cause a neurodevelopmental disorder termed MYCBP2-related Developmental delay with Corpus callosum Defects (MDCD) (Al-Abdi et al., 2022). MDCD features defective neuronal connectivity including a hypoplastic or absent corpus callosum, neurobehavioral deficits including intellectual disability and epilepsy, and abnormal craniofacial development. This constellation of comorbidities in MDCD closely resembles some of the phenotypes observed in mice with deficient EPHB2 signalling. Given our finding that MYCBP2 loss reduces EPHB2 levels and influences Eph receptor effects on axon development, the question of whether EPHB2 expression levels are normal in MYCBP2 patients remains pertinent.

Revised version: Importantly, the biomedical relevance of our findings are heightened by a recent study that identified genetic variants in MYCBP2, which cause a neurodevelopmental disorder termed MYCBP2-related Developmental delay with Corpus callosum Defects (MDCD) (AlAbdi et al., 2023). MDCD features defective neuronal connectivity including a hypoplastic or absent corpus callosum, neurobehavioral deficits including intellectual disability and epilepsy, and abnormal craniofacial development. This constellation of comorbidities in MDCD closely resembles some of the phenotypes observed in mice with deficient EPHB2 signalling. Given our finding that MYCBP2 loss reduces EPHB2 levels and influences Eph receptor effects on axon development, the question of whether EPHB2 expression levels are normal in MYCBP2 patients remains pertinent.

In conclusion, our study has revealed numerous biochemical and genetic links between MYCBP2 and EPHB2. Our findings indicate that the MYCBP2/FBXO45 complex protects EPHB2 from degradation, and these are functionally integrated signalling players with an evolutionarily conserved role in axonal development. Future studies will be needed to address how the EPHB2-MYCBP2 interaction affects nervous system development in mammals in vivo and to identify further regulators of EPHB2 degradation. Additionally, another idea worthy of closer examination in the future is the possibility that MYCBP2 signalling could provide routes through which EPHB2-initiated signals access numerous fundamental cellular functions.